# Analiza danych w języku Python - ćwiczenia laboratoryjne 2024/2025

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika. Wyjątek stanowi ostatni z projektów, gdzie po konsultacji z prowadzącym zajęcia możliwa jest większa swoboda w doborze pakietów.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają na celu wyłapać podstawowe błędy. Przejście przez kod wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem). Rozmiar poziomy figur nie powinien przekraczać 20 cali.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (nazwa zespołu oraz imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespołu:
Członkowie:

---

## Zestaw zadań 4: Wstęp do analizy danych przestrzennych (Notatnik 1/2)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
import geopandas as gpd
import os
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Zadanie 1: Import i przygotowanie danych (9 pkt)

Wykorzystując funkcję `geopandas.read_file()` zaimportuj do notatnika dane z załączonych do niego plików i przygotuj na ich podstawie cztery tablice zawierające wybrane dane dotyczące dwóch rzek - Wisły i Odry, miast, powiatów oraz województw.

Wyświetl przygotowane tablice i zapisz je do pliu `data.gpkg` na warstwy o nazwach odpowiadające nazwom zmiennych.

Uwagi do wykonania zadania:
 - Pliki zawierają więcej danych niż wynikowe tablice, w których nie powinny znaleźć się zbędne dane.
 - Poprawny import polskich nazw może wymagać ustawienia odpowiedniej wartości parametru `encoding` (w zależności od pliku utf-8 lub 1250) podczas wywoływania funkcji `geopandas.read_file()`.
 - Dane są zapisane z wykorzystaniem różnych odwzorowań kartograficznych. Przed rozpoczęciem analizy należy je ze sobą uzgodnić. Korzystając z własności obiektu GeoDataFrame `geopandas.GeoDataFrame.crs` sprawdź odwzorowanie kartograficzne, w którym zapisane są dane. Jeżeli odwzorowanie jest inne niż epsg:2180 dokonaj jego konwersji z wykorzystaniem funkcji `geopandas.GeoDataFrame.to_crs()`.

#### a)  Rzeki

Tablica `rivers` zawierająca wyłącznie następujące elementy:</br>

> Indeks:</br>
> Domyślny indeksowanie tablicy

> Kolumny: Nazwa, geometry</br>
> Kolumny zawierające nazwę oraz geometrię obiektów

W tablicy zawrzyj wyłącznie informacje dotyczące dwóch rzek - Wisły i Odry.

<!-- grade_id: cell-a3f09fcd0c1b8b0e, points: 0, grade: false, solution: true -->
```python
df=geopadas.read_file('
```

<!-- grade_id: cell-fa72e0d0e0547422, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(rivers) == gpd.geodataframe.GeoDataFrame
assert rivers.shape == (2, 2)
assert list(rivers.columns) == ['Nazwa', 'geometry']
assert rivers.crs == "epsg:2180"
```

#### b) Miasta

Tablica `cities` zawierająca wyłącznie następujące elementy:</br>

> Indeks:</br>
> Domyślny indeksowanie tablicy

> Kolumny: Nazwa, geometry</br>
> Kolumny zawierające nazwę oraz geometrię obiektów

W tablicy zawrzyj wyłącznie informacje dotyczące miast jako całości, pomiń wsie, dzielnice itp.

<!-- grade_id: cell-6c1e035b246fbf9f, points: 0, grade: false, solution: true -->
```python
# Miasta
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3a5d54392278957c, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(cities) == gpd.geodataframe.GeoDataFrame
assert cities.shape == (964, 2)
assert list(cities.columns) == ['Nazwa', 'geometry']
assert cities.crs == "epsg:2180"
```

#### c) Powiaty

Tablica `districts` zawierająca wyłącznie następujące elementy:</br>

> Indeks:</br>
> Domyślny indeksowanie tablicy

> Kolumny: Nazwa, geometry</br>
> Kolumny zawierające nazwę oraz geometrię obiektów

<!-- grade_id: cell-bb817f505b502999, points: 0, grade: false, solution: true -->
```python
# Powiaty
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-064888548be59294, points: 1, grade: true, solution: false -->
```python
# Komórka testowa
assert type(districts) == gpd.geodataframe.GeoDataFrame
assert districts.shape == (380, 2)
assert list(districts.columns) == ['Nazwa', 'geometry']
assert districts.crs == "epsg:2180"
```

#### d) Województwa

Tablica `voivodeships` zawierająca wyłącznie następujące elementy:</br>

> Indeks:</br>
> Domyślny indeksowanie tablicy

> Kolumny: komplet kolumn z pliku `voivodeships_dane.csv` oraz przygotowana osobno kolumna geometry</br>
> Kolumny zawierające nazwę, dodatkowe inforamacje o obiektach oraz ich geometrię

Plik `voivodeships_dane.csv` nie zawiera geometrii obiektów. W związku z tym należy go wczytać z wykorzystaniem funkcji `pandas.read_csv()`, a geometrię obiektów przygotować osobno, korzystając z funkcji `geopandas.GeoDataFrame.dissolve()`, w oparciu o dane zawarte w pliku `Powiaty.zip`.

<!-- grade_id: cell-2cd2490da8292fc9, points: 0, grade: false, solution: true -->
```python
# Województwa
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c3ad9ba3706dcfa9, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert type(voivodeships) == gpd.geodataframe.GeoDataFrame
assert voivodeships.shape == (16, 13)
assert set(list(voivodeships.columns)) == set(['Nazwa',
                                             'geometry',
                                             'Powierzchnia ogółem [ha]',
                                             'Produkcja energii elektrycznej [GWh]',
                                             'Stosunek produkcji energii elektrycznej do zużycia energii elektrycznej [%]',
                                             'Udział energii odnawialnej w produkcji energii elektrycznej [%]',
                                             'Udział powierzchni gruntów pod wodami w powierzchni ogółem [%]',
                                             'Udział powierzchni nieużytków w powierzchni ogółem [%]',
                                             'Udział powierzchni terenów mieszkaniowych w powierzchni ogółem [%]',
                                             'Udział powierzchni terenów przemysłowych w powierzchni ogółem [%]',
                                             'Udział powierzchni terenów rekreacji i wypoczynku w powierzchni ogółem [%]',
                                             'Udział powierzchni użytków rolnych w powierzchni ogółem [%]',
                                             'Zużycie energii elektrycznej [GWh]'])
assert voivodeships.crs == "epsg:2180"
```

#### e) Zapisanie danych

Korzystając z funkcji `geopandas.GeoDataFrame.to_file()` zapisz dane z przygotowanych tablic do pliku `data.gpkg` na osobne warstwy o nazwach identycznych z nazwami zapisywanych zmiennych.

Znajdujący się w komórce kod usuwa już istniejący plik `data.gpkg`. Zapobiega to nadpisywaniu danych do już istniejącego pliku w kolejnych wykonaniach notatnika.

<!-- grade_id: cell-d6d61e494f405909, points: 0, grade: false, solution: true -->
```python
# Zapisanie danych do pliku
if os.path.exists("data.gpkg"):
    os.remove("data.gpkg")
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ed5493dd66a0d5ce, points: 1, grade: true, solution: false -->
```python
# Komórka testowa
test_results = gpd.list_layers("data.gpkg")
assert test_results.shape == (4, 2)
assert np.all(set(test_results['name'].values) == set(['rivers', 'cities', 'districts', 'voivodeships']))
assert np.all(set(test_results['geometry_type'].values) == set(['MultiLineString', 'Point', 'MultiPolygon', 'MultiPolygon']))
```

### Zadanie 2: Wizualizacja danych wejściowych (2 pkt)

Korzystając z funkcji pakietu matplotlib oraz funkcji `geopandas.GeoDataFrame.plot()` zwizualizuj na pojedynczej mapie komplet przygotowanych danych (geometrie obiektów). Przygotuj wykres tak, aby:
 - powiaty posiadały wypełnienie w wybranym kolorze oraz granice oznaczone cienką linią w bardziej intensywnej wersji koloru ich wypełnienia,
 - granice województw były oznaczne grubszą linią w bardziej intensywnej wersji koloru granic powiatów,
 - rzeki i miasta były oznaczone kolorami wyróżniającymi się na tle oznaczeń powiatów i województw.

Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Ustaw parametry wykresu oraz wizualizowanych obiektów (wielkość figury, kolory, grubość linii, wielkość punktów itp.) tak aby wygenerowana mapa była przejrzysta i czytelna.

<!-- grade_id: cell-7cf2575ea3586fc8, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Zadanie 3: Mierzenie długości, pola powierzchni i odległości (7 pkt)

#### a) Długość Wisły
Korzystając z własności `geopandas.GeoSeries.length` zmierz długość Wisły w metrach oraz:
 - zapisz wynik do zmiennej `length`,
 - korzystając z funkcji `print()` wyświetl odpowiedź (pełnym zdaniem) na pytanie o długość Wisły podając odpowiedź w kilometrach z dokładnością do 3 miejsc po przecinku.

<!-- grade_id: cell-15ff549b14ca0466, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-24758c9f3da5ccde, points: 1, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(length, 1019757.7991167122)
```

#### b) Pola powierzchni województw
Korzystając z własności `geopandas.GeoSeries.area` zmierz powierzchnię województw w metrach kwadratowych oraz:
 - zapisz wyniki do tablicy `voivodeships` w kolumnie o nazwie "Pole", 
 - oblicz pole powierzchni Polski w metrach kwadartowych i zapisz wynik do zmiennej `area`,
 - korzystając z funkcji `print()` wyświetl odpowiedź na pytanie (pełnym zdaniem) o powierzchnię Polski podając odpowiedź w kilometrach kwadratowych z dokładnością do 3 miejsc po przecinku,
 - przygotuj wizualiację, na której kolor przypisany województwom jest powiązany z ich powierzchnią.

Wizualizacja powinna zawierać:
- województwa, wypełnione kolorem powiązanym z ich powierzchnią (ciągła skala kolorów),

Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Ustaw parametry wykresu oraz wizualizowanych obiektów (wielkość figury, kolory, grubość linii, wielkość punktów itp.) tak aby wygenerowana mapa była przejrzysta i czytelna.

<!-- grade_id: cell-eb301c3344f17b1e, points: 0, grade: false, solution: true -->
```python
# Obliczenia
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-54e8fb04359f397e, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert np.all(np.isclose(np.sort(voivodeships["Pole"].values), np.sort(np.array([1.99361725e+10, 1.79475269e+10, 2.51340028e+10, 1.39898166e+10,
                                                               3.55292977e+10, 1.51660383e+10, 9.39998646e+09, 1.78437052e+10,
                                                               2.01931654e+10, 1.83069651e+10, 2.41505782e+10, 2.97972921e+10,
                                                               2.29033672e+10, 1.81944352e+10, 1.23170381e+10, 1.16972569e+10]))))
assert np.isclose(voivodeships[voivodeships["Nazwa"]=="małopolskie"]["Pole"].values[0], 15166038294.192154)
assert np.isclose(voivodeships[voivodeships["Nazwa"]=="pomorskie"]["Pole"].values[0], 18306965109.079857)
assert np.isclose(area, 312506644577.3843)
```

<!-- grade_id: cell-ab1428f12655827b, points: 1, grade: true, solution: true -->
```python
# Wizualizacja
# YOUR CODE HERE
raise NotImplementedError()
```

#### c) Odległość miast od Krakowa
Korzystając z funkcji `geopandas.GeoSeries.distance()` zmierz odległość poszczególnych miast od Krakowa w metrach oraz:
 - zapisz wyniki do tablicy `cities` w kolumnie o nazwie "Odległość od Krakowa", 
 - oblicz w jakiej odległości od Krakowa w kilometrach leży Poznań i zapisz wynik do zmiennej `distance`,
 - korzystając z funkcji `print()` wyświetl odpowiedź na pytanie (pełnym zdaniem) w jakiej odległości wyrażonej w kilometrach od Krakowa leży Poznań,
 - przygotuj wizualiację, na której kolor przypisany miastom jest powiązany z ich odległością od Krakowa. 

Wizualizacja powinna zawierać:
- kontury województw, bez wypełnienia,
- lokalicaje miast, których kolor jest powiązany z ich odległością od Krakowa (ciągła skala kolorów).

Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Ustaw parametry wykresu oraz wizualizowanych obiektów (wielkość figury, kolory, grubość linii, wielkość punktów itp.) tak aby wygenerowana mapa była przejrzysta i czytelna.

<!-- grade_id: cell-ce6d9808f4193a3c, points: 0, grade: false, solution: true -->
```python
# Obliczenia
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ea3373eed57cd3e3, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(cities["Odległość od Krakowa"].min(), 0.0)
assert np.isclose(cities["Odległość od Krakowa"].max(), 579020.726082644)
assert np.isclose(cities["Odległość od Krakowa"].mean(), 271127.08178774477)
assert np.isclose(distance, 334.60495145299524)
```

<!-- grade_id: cell-8d841cc379d33cb6, points: 1, grade: true, solution: true -->
```python
# Wizualizacja
# YOUR CODE HERE
raise NotImplementedError()
```